# CSCI 3155 : Homework 6



__Name__: Ian McGregor

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
passed: (points: Int)Unit
```

## Problem 1 (20 points): Map and filter with fold


```scala
sealed trait MyList[+A]
case object Empty extends MyList[Nothing]
case class Cons[A]( x : A, xs : MyList[A]) extends MyList[A]

def fold[A,B](f : (B,A) => B, acc : B, list : MyList[A]) : B = 
    list match {
        case Empty => acc
        case Cons(h,t) => fold(f, f(acc,h), t)
    }
```

Output:
```
defined trait MyList
defined object Empty
defined class Cons
fold: [A, B](f: (B, A) => B, acc: B, list: MyList[A])B
```

Here are the ``snoc`` and ``concat`` list utility functions. These might be useful for the rest of problem 1.

The function ``snoc`` adds an element to the end of a list : 

``snoc( [1,2,3], 4 ) = [1,2,3,4]``


The function ``concat`` concatenates two lists : 

``concat( [1,2,3], [4,5,6] ) = [1,2,3,4,5,6]``


```scala
def snoc[A]( list : MyList[A], a : A) : MyList[A]= 
    list match {
        case Empty => Cons( a, Empty)
        case Cons( h, t) => Cons( h, snoc(t,a))
    }

snoc( Cons(1, Cons(2, Cons(3, Empty))), 4)
snoc( snoc( snoc( snoc(Empty, 1), 2), 3), 4)

def concat[A]( list1: MyList[A], list2 : MyList[A]) : MyList[A] =
    list2 match {
        case Empty => list1
        case Cons( h, t) => concat( snoc(list1,h), t)
    }

concat( Cons(1, Cons(2, Cons(3, Empty))), Cons(4, Cons(5, Cons(6, Empty))))
```

Output:
```
Cons(1,Cons(2,Cons(3,Cons(4,Cons(5,Cons(6,Empty))))))
```

### Problem 1a

Write the ``map`` function for the ``MyList`` type using ``fold``.

<!-- grade_id: cell-9e11cd3de61ea999, points: 0, grade: false, solution: true -->
```scala
def map1a[A,B]( f : A => B, list : MyList[A]) : MyList[B] = {
    fold( (x : MyList[B] , y : A ) => snoc (x , f(y)) , Empty, list)
}
```

Output:
```
map1a: [A, B](f: A => B, list: MyList[A])MyList[B]
```

<!-- grade_id: cell-515acf15795db239, points: 10, grade: true, solution: false -->
```scala
val list1 = Cons(1, Cons(2, Cons(3, Cons(4, Cons(5, Empty)))))
val list2 = Cons(2, Cons(4, Cons(6, Cons(8, Cons(10, Empty)))))
val list3 = Cons("1", Cons("2", Cons("3", Cons("4", Cons("5", Empty)))))

assert( map1a( (x : Int) => 2 * x, Empty) == Empty , "test 1")
assert( map1a[Int,Int]( x => 2 * x, list1) == list2 , "test 2")
assert( map1a[Int,String]( x => x.toString, list1) == list3, "test 3")
assert( map1a( (x : String) => x.toInt * 2, list3) == list2, "test 4")
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
null
```

### Problem 1b

Write the ``filter`` function for the ``MyList`` type using ``fold``.

<!-- grade_id: cell-ba1618e5d398ed3d, points: 0, grade: false, solution: true -->
```scala
def filter1b[A](f : A => Boolean, list : MyList[A]) : MyList[A] = 
    fold( (x : MyList[A] , y : A ) => if (f(y) == true){snoc(x,y)}else x, Empty, list)
```

<!-- grade_id: cell-3de5bb68ed294aaa, points: 10, grade: true, solution: false -->
```scala
val list1 = Cons(1, Cons(2, Cons(3, Cons(4, Cons(5, Empty)))))
val list2 = Cons(2, Cons(4, Empty))
val list3 = Cons(1, Cons(3, Cons(5, Empty)))
val list4 = Cons("one", Cons("two", Cons("three", Cons("four", Cons("five", Empty)))))
val list5 = Cons("three", Cons("four", Cons("five", Empty)))

assert( filter1b( (x : Int) => x % 2 == 0, Empty) == Empty, "test 1")
assert( filter1b( (x : Int) => x % 2 == 0, list1) == list2, "test 2")
assert( filter1b[Int]( x => x % 2 == 1, list1) == list3, "test 3")
assert( filter1b( (x : String) => x.length > 3, list4) == list5, "test 4")
passed(10)
```

## Problem 2 (20 points): Map, filter, fold, and lambda practice.

Here are some practice problems using map filter and fold over scala lists. Write the functions as a lambda that takes a list as an argument and then applies one or more of the map filter or fold functions to it to get the result

```scala
// Example 
// get the sum of a list

val sum : List[Int] => Int =
    list =>
        list.foldLeft(0)( (x,y) => x + y)

val test = List(1,2,3,4,5)
assert( sum(test) == 15)
```

<!-- grade_id: cell-8cf280e696124d65, points: 0, grade: false, solution: true -->
```scala
// Problem 2a
// sum all the even numbers in a list
val sumeven : List[Int] => Int = 
    list => 
        list.filter((f : Int) => f % 2 == 0).foldLeft(0)( (x,y) => x + y)
```

<!-- grade_id: cell-7c23b49a903043df, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1,2,3,4,5)
val test2 = List(108,342,6543,648,217,214)
assert(sumeven(test1) == 6)
assert(sumeven(test2) == 1312)
passed(2)
```

<!-- grade_id: cell-03dcf2a103cac464, points: 0, grade: false, solution: true -->
```scala
// Problem 2b
// find the max of a list of positive numbers 

val max : List[Int] => Int = 
    list => {
        list.foldLeft(0)( (x,y) => if(x<y) y else x)
    }
```

<!-- grade_id: cell-fc94abed35a74dc1, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1,2,3,4,5)
val test2 = List(345,234,5,76,354,2415,7645,974)
assert(max(test1) == 5)
assert(max(test2) == 7645)
passed(2)
```

<!-- grade_id: cell-18ce05d1bbf6f67b, points: 0, grade: false, solution: true -->
```scala
// Problem 2c
// replace all odd numbers in a list with their double
// def map[A,B](function : A => B, list : List[A]) : List[B] = {
//     list match {
//         case Empty => Empty
//         case Cons(h, t) => Cons(function(h), map(function, t))
//     }
// }

val replaceodd : List[Int] => List[Int] = 
    list => {
        list.map( (x : Int) => 
                 if(x%2 == 1){
                     x * 2
                 }
                 else{
                     x
                 })
    }
```

<!-- grade_id: cell-0400375448b35757, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1,2,3,4,5)
val test2 = List(17,16,42,25,6,92,18,49)
assert(replaceodd(test1) == List(2,2,6,4,10))
assert(replaceodd(test2) == List(34, 16, 42, 50, 6, 92, 18, 98))
passed(2)
```

<!-- grade_id: cell-443681a4843bf4a9, points: 0, grade: false, solution: true -->
```scala
// Problem 2d
// get the multiplicative inverse of a list of numbers, but drop any zeros.
// aka reciprocal aka inverse aka flip it and reverse it as long as its not zer0000

val inverse : List[Double] => List[Double] = 
    list => 
        list.filter((x : Double) => x!=0).map((x:Double) => 1/x)

```

<!-- grade_id: cell-f6e17fcb6c21344b, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1.0,0.0,1.0/3.0,0.0,5.0)
val test2 = List(0.25, 0.0, 4.0)
assert(inverse(test1) == List(1.0, 3.0, 1.0/5))
assert(inverse(test2) == List(4.0, 0.25))
passed(2)
```

<!-- grade_id: cell-80158c8add776398, points: 0, grade: false, solution: true -->
```scala
// Problem 2e
// reverse a list

val reverse : List[Int] => List[Int] = {
    list => 
        list.foldLeft(List[Int]())((a,b) => b :: a)
}
val test2 = List(36, 29, 71, 61, 21, 66, 47, 27, 98, 60)
reverse(test2)
```

<!-- grade_id: cell-4e71ae8b353b4bef, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1,2,3,4,5)
val test2 = List(36, 29, 71, 61, 21, 66, 47, 27, 98, 60)
assert(reverse(test1) == List(5, 4, 3, 2, 1))
assert(reverse(test2) == List(60, 98, 27, 47, 66, 21, 61, 71, 29, 36))
passed(2)
```

<!-- grade_id: cell-9306547a582cc03c, points: 0, grade: false, solution: true -->
```scala
// Problem 2f
// get the square bracket string representation of a list

val printlist : List[Int] => String = {
       list => 
            "[" + list.fold("")((x, y) => if(x!=""){x.toString + ", " + y.toString }else{ y.toString}) + "]"
}
```

<!-- grade_id: cell-732b2e0f451e4bff, points: 2, grade: true, solution: false -->
```scala
val test1 = List(1,2,3,4,5)
val test2 = List()
val test3 = List(23, 7, 494, 2884, 17, 6, 38)
assert(printlist(test1) == "[1, 2, 3, 4, 5]")
assert(printlist(test2) == "[]")
assert(printlist(test3) == "[23, 7, 494, 2884, 17, 6, 38]")
passed(2)
```

<!-- grade_id: cell-5345574ae122b1f6, points: 0, grade: false, solution: true -->
```scala
// Problem 2g
// count all points within radius 4 of (0,0)
val circle : List[(Double,Double)] => Int =
    list => list.filter((x) => x._1 * x._1 + x._2 * x._2 < 16).foldLeft(0)((x,y) => x+1)
```

<!-- grade_id: cell-88066ff345373719, points: 2, grade: true, solution: false -->
```scala
val test1 = List( (0.0,0.0), (1.0,2.0), (3.0,4.0), (5.0,6.0))
val test2 = List( (2.36 , -3.10), (-3.22 , 1.33), (4.23 , 2.99),
                  (1.42 , -3.10), (4.43 , -1.95), (-0.78 , -4.81), 
                  (4.78 , 4.58), (-4.78 , -3.81), (2.62 , 0.04) ,(0.40 , -1.01))
assert(circle(test1) == 2)
assert(circle(test2) == 5)
passed(2)
```

<!-- grade_id: cell-ad40a3f50758f563, points: 0, grade: false, solution: true -->
```scala
// Problem 2h
// calculate calories are in a meal

sealed trait Food
case object Water extends Food
case class IceCream(flavor : String, calories : Int) extends Food
case class Salad(calories : Int) extends Food
case class Sandwich(layers : Int, calories_per_layer : Int) extends Food

val caloriecount : List[Food] => Int =
    list =>
        list.map ( x => x match {
            case Water => 0
            case IceCream(sabor, cals) => cals
            case Salad(cals) => cals
            case Sandwich(lay, cals) => lay * cals
        }).foldLeft(0)((x,y) => x + y)
```

<!-- grade_id: cell-898c69a568bc7fe4, points: 2, grade: true, solution: false -->
```scala
val meal1 = List( Water, Salad(10))
val meal2 = List( Sandwich(3, 100), IceCream("chocolate", 200), Water, IceCream("vanilla", 300))
assert(caloriecount(meal1) == 10)
assert(caloriecount(meal2) == 800)
passed(2)
```

<!-- grade_id: cell-5899048198f34809, points: 0, grade: false, solution: true -->
```scala
// Problem 2i
// flatten a list of lists into a single list of elements

val flatten : List[List[Int]] => List[Int] = 
    list =>
        
```

<!-- grade_id: cell-836405fa9ec40614, points: 2, grade: true, solution: false -->
```scala
val test1 = List( List(1), List(2,2), List(3,3,3))
val test2 = List( List(43,1), List[Int](), List(9,5,10), List(11))
assert(flatten(test1) == List(1, 2, 2, 3, 3, 3))
assert(flatten(test2) == List(43, 1, 9, 5, 10, 11))
passed(2)
```

<!-- grade_id: cell-ca3a805fb51ff2ce, points: 0, grade: false, solution: true -->
```scala
// Problem 2j
// create a map from a list, without adding existing keys

val makemap : List[(Int,String)] => Map[Int, String] = 
    list => 
        list 
```

<!-- grade_id: cell-d863dc73ee02edcb, points: 2, grade: true, solution: false -->
```scala
val test1 = List( (1,"one"), (2, "two"), (3, "three"), (4,"four"), (5, "five"))
val test2 = List( (10, "ten"), (2, "2"), (10, "10"), (2,"two"), (10, "0xA"))
assert(makemap(test1) == Map(1 -> "one", 2 -> "two", 3 -> "three", 4 -> "four",  5 -> "five"))
assert(makemap(test2) == Map(10 -> "ten", 2 -> "2"))
passed(2)
```

## Problem 3 (20 points): Dictionaries

We define the datatype $\texttt{Dictionary}$ using the definition below:

$$
\begin{aligned}
\textbf{Dictionary}\ A\ B\ ::&= \text{EmptyDict} \\
&\mid\ \text{Entry}\ A\ B\ (\textbf{Dictionary}\ A\ B)
\end{aligned}
$$

This will create dictionaries of key-value pairs which will be useful for implementing environments. The definition of this type in Scala is given below:

```scala
sealed trait Dictionary[+A, +B] 
case object EmptyDict extends Dictionary[Nothing, Nothing]
case class Entry[A,B](key : A, value : B, xs : Dictionary[A, B]) extends Dictionary[A, B]
```

### Problem 3a

Write a function called `contains` which returns `True` if the Dictionary contains a given key and `False` if it does not.

$$
\texttt{contains} :  \texttt{Dict}\ a\ b \rightarrow a \rightarrow Boolean
$$

<!-- grade_id: cell-19e9eb529edb079b, points: 0, grade: false, solution: true -->
```scala
def contains[A,B](dict : Dictionary[A,B], k : A) : Boolean = 
    dict match{
        case EmptyDict => false
        case Entry(key, valueb, diction) => (key == k) match {
            case true => true
            case false => contains(diction, k)
        }
    }
```

<!-- grade_id: cell-16ed2600c99a7834, points: 5, grade: true, solution: false -->
```scala
assert(contains(Entry(5, true, EmptyDict), 5), "test 1")
assert(contains(Entry(8, true, Entry(5, true, EmptyDict)), 5),"test 2")
assert(!contains(Entry(5, true, EmptyDict), 6), "test 3")
assert(!contains(EmptyDict, 7))
passed(5)
```

### Problem 3b

Write a function ``lookup`` that retrieves a value that corresponds to a key in a dictionary. If the key is not in the map then return None.

$$
\texttt{lookup} :  \texttt{Dict}\ a\ b \rightarrow a \rightarrow \texttt{Maybe}\ b
$$

You will need to implement the ``Maybe`` type.

<!-- grade_id: cell-f36e0a24a3d2c697, points: 0, grade: false, solution: true -->
```scala
sealed trait Maybe[+A]
case object None extends Maybe[Nothing]
case class Some[A] (x : A) extends Maybe[A]

def lookup[A,B] (dict : Dictionary[A,B], k : A) : Maybe[B] = 
    dict match{
        case EmptyDict => None
        case Entry(x,y,diction) => x == k match{
            case true => Some(y)
            case false => lookup(diction,k)
        }
    }
```

<!-- grade_id: cell-ce50b930762797e4, points: 5, grade: true, solution: false -->
```scala
assert(lookup(Entry(5, "ham", EmptyDict), 5) == Some("ham"))
assert(lookup(Entry(8, 7, Entry(5, 3, EmptyDict)), 5) == Some(3))
assert(lookup(Entry(5, true, EmptyDict), 6) == None)
assert(lookup(EmptyDict, 7) == None)
passed(5)
```

### Problem 3d

Create the following dictionary as a scala value using this definition:
```
{
    five: True,
    six: False,
    seven: False
}
```
Place it in a variable named "`my_dict`"

<!-- grade_id: cell-328e9f8e1c4f0c64, points: 0, grade: false, solution: true -->
```scala
val my_dict = Entry(5, true, Entry(6 , false, Entry(7, false, EmptyDict)))
```

<!-- grade_id: cell-60b14443b5138a7d, points: 2, grade: true, solution: false -->
```scala
assert(lookup(my_dict, 5) == Some(true))
assert(lookup(my_dict, 6) == Some(false))
assert(lookup(my_dict, 7) == Some(false))
assert(lookup(my_dict, 8) == None)
passed(2)
```

### Problem 3e
For the following, use the super secret `secret_dict` variable defined below. Use the functions you just defined to complete the following:

1. Look up the value mapped by key `1` and place it in variable `a` 
2. Look up the value mapped by key `3` and place it in variable `b`
3. Look up the value mapped by the key `100` and place it in the variable `c`
4. Check whether the key `-3` is in the map and place True or False in variable `d`
5. Check whether the key `-1` is in the map and place True or False in variable `e`


**Note**: For 1,2,3 remember to extract it from the maybe type, set the variable to the string `"does not exist"` if you get `None` instead of `Some(x)`.

```scala
val secret_dict = Entry(3, "abc",
                        Entry(-1, "xyz",
                              Entry(1, "secret", EmptyDict)))
```

<!-- grade_id: cell-18d5d9a267c6e266, points: 0, grade: false, solution: true -->
```scala
val a = lookup(secret_dict, 1) match {
    case Some(x) => x
    case None => "does not exist"
}
val b = lookup(secret_dict, 3) match {
    case Some(x) => x
    case None => "does not exist"
}
val c = lookup(secret_dict, 100) match {
    case Some(x) => x
    case None => "does not exist"
}
val d = contains(secret_dict, -3)
val e = contains(secret_dict, -1)
```

<!-- grade_id: cell-27a2f9fa23da9abb, points: 3, grade: true, solution: false -->
```scala
// Hidden tests (3 pts)
```

### Problem 3f

Write a ``eval`` function for this subset of lettuce using the Dictionaries we have defined to store the environment.


```scala
sealed trait Expr
case class Const(x : Integer) extends Expr
case class Ident(x : String) extends Expr
case class Plus(x : Expr, y : Expr) extends Expr
case class Let(id : String, x : Expr, y : Expr) extends Expr
```

<!-- grade_id: cell-b0178d64c649d3b3, points: 0, grade: false, solution: true -->
```scala
def eval( expr : Expr, env : Dictionary[String, Int]) : Int = 
    expr match { 
        case Const(a) => a
        case Ident(b) => lookup(env,b) match {
            case Some(x) => x
            case None => throw new Exception(s"not in environment")
            }
        case Plus(c,d) => eval(c,env) + eval(d,env)
        case Let(e, f, g) => eval(g, Entry(e,eval(f,env),env))
        
    }
```

<!-- grade_id: cell-f1840face810867c, points: 5, grade: true, solution: false -->
```scala
val expr1 = Plus( Const(3), Const(4))
val expr2 = Let("x", Const(10),Plus( Ident("x"), Const(1)))
val expr3 = Let("y", Const(20), Let("x", Const(10), Plus( Ident("x"), Ident("y"))))
val expr4 = Let("x", Let("y", Const(10), Plus( Ident("y"), Const(5))), Plus(Ident("x"), Const(5)))

assert(eval(expr1, EmptyDict) == 7)
assert(eval(expr2, EmptyDict) == 11)
assert(eval(expr3, EmptyDict) == 30)
assert(eval(expr4, EmptyDict) == 20)
passed(5)
```

```scala

```