```scala

```

# CSCI 3155 : Assignment 8

__Name__:

Ian McGregor

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
passed: (points: Int)Unit
```

## Problem 1 : Type Checking (100 points)

Write a function `typecheck` that typechecks a Lettuce program.

$$\texttt{typecheck} : Program \rightarrow Type$$ 

This function should take as input a lettuce _Program_, and return what _Type_ it is. If it doesn't typecheck it should return a special value _TypeError_. Use the Lettuce typing notes on moodle for details on how typechecking lettuce works.

The tests for this assigment are hidden, so you will need to write your own tests! Note that you are encouraged to collaborate and share your test ideas on Piazza. To get full points, your implementation needs to match the specification from the Lettuce typing notes on moodle.

```scala
sealed trait Type
case object NumType extends Type
case object BoolType extends Type
case class FunType( ty_a : Type, ty_r : Type) extends Type
case object TypeError extends Type

sealed trait Expr
case class  Const(n : Int) extends Expr
case class  Ident(s : String) extends Expr
case object True extends Expr
case object False extends Expr
case class  Plus( e1 : Expr, e2 : Expr) extends Expr
case class  Minus( e1 : Expr, e2 : Expr) extends Expr
case class  Mult( e1 : Expr, e2 : Expr) extends Expr
case class  Neg( e1 : Expr) extends Expr
case class  And( e1 : Expr, e2 : Expr) extends Expr
case class  Or( e1 : Expr, e2 : Expr) extends Expr
case class  Geq( e1 : Expr, e2 : Expr) extends Expr
case class  Eq( e1 : Expr, e2 : Expr) extends Expr
case class  ITE( ec : Expr, et : Expr, ef : Expr) extends Expr
case class  Let( name : String, ty_d : Type, ed : Expr, ei : Expr) extends Expr
case class  FunDef( arg : String, ty_a : Type, eb : Expr) extends Expr
case class  FunCall( ef : Expr, ea : Expr) extends Expr
case class  LetRec( f : String, ty_f : Type, x : String, ty_x : Type, eb : Expr, ei : Expr) extends Expr

sealed trait Program
case class Main(p : Expr) extends Program
```

Output:
```
defined trait Type
defined object NumType
defined object BoolType
defined class FunType
defined object TypeError
defined trait Expr
defined class Const
defined class Ident
defined object True
defined object False
defined class Plus
defined class Minus
defined class Mult
defined class Neg
defined class And
defined class Or
defined class Geq
defined class Eq
defined class ITE
defined class Let
defined class FunDef
defined class FunCall
defined class LetRec
defined trait Program
defined class Main
```

<!-- grade_id: cell-f8cd57266fd56cf2, points: 0, grade: false, solution: true -->
```scala
type Environment = 
    Map[String, Type]

def TC (input : Expr, gamma : Environment) : Type = {
    input match {
        case Const(n) => n match {
            case n : Int => NumType
            }
        case Ident(s) => if(gamma.contains(s)) gamma(s) else TypeError
        case True => BoolType
        case False => BoolType
        case Plus(e1,e2) => (TC(e1,gamma) , TC(e2, gamma)) match {
            case (NumType, NumType) => NumType
            case (_,_) => TypeError
        }
        case Minus(e1,e2) => (TC(e1,gamma) , TC(e2, gamma)) match {
            case (NumType, NumType) => NumType
            case (_,_) => TypeError
        }
        case Mult(e1,e2) => (TC(e1,gamma) , TC(e2, gamma)) match {
            case (NumType, NumType) => NumType
            case (_,_) => TypeError
        }
        case Neg(e1) => TC(e1,gamma) match {
            case TypeError => TypeError
            case NumType => NumType
            case BoolType => BoolType
            case FunType(a,b) => FunType(a,b)
        }
        case And(e1,e2) => (TC(e1,gamma) , TC(e2, gamma)) match { 
            case (BoolType,BoolType) => BoolType
            case (_,_) => TypeError
        }
        case Or(e1,e2) => (TC(e1,gamma) , TC(e2, gamma)) match { 
            case (BoolType,BoolType) => BoolType
            case (_,_) => TypeError
        }
        case Geq(e1,e2) => (TC(e1,gamma) , TC(e2, gamma)) match {
            case (NumType, NumType) => BoolType
            case (_,_) => TypeError
        }
        case Eq(e1,e2) => (TC(e1,gamma) , TC(e2, gamma)) match {
            case (NumType, NumType) => BoolType
            case (BoolType, BoolType) => BoolType
            case (_,_) => TypeError
        }
        case ITE(ec, et , ef) => TC(ec, gamma) match {
            case TypeError => TypeError
            case NumType => TypeError
            case FunType(_,_) => TypeError
            case BoolType => (TC(et,gamma)) match {
                case TypeError => TypeError
                case et => if( et == TC(ef,gamma)){
                    et
                }else{TypeError}
            }
        }
        case Let(name, ty_d, ed, ei) => TC(ed,gamma) match {
            case TypeError => TypeError
            case ty_d => {
                val nugamma = gamma + (name -> TC(ed,gamma))
                TC(ei, nugamma)
            }
        }
        case FunDef(arg, ty_a, eb) => {
            val nugamma = gamma + (arg -> ty_a)
            TC(eb, nugamma) match {
                case TypeError => TypeError
                case tr => {
                    FunType(nugamma(arg), tr)
                }
            }
        }
        case FunCall(ef, ea) => TC(ef, gamma) match {
            case TypeError => TypeError
            case NumType => TypeError
            case BoolType => TypeError
            case FunType(ta, tr) => TC(ea, gamma) match {
                case ta => {
                    tr
                }
            }   
        }
        case LetRec(f_name, ty_f, arg_name, ty_arg, eb, ei) => {
            val nugamma = gamma + (arg_name -> ty_arg) + (f_name -> ty_f)
            TC(ei, nugamma) match {
                case TypeError => TypeError
                case ti => TC(eb,nugamma)
            }
        }
    }
}
def typecheck (program : Program) : Type = {
    val emptyenv : Environment = Map()
    program match {
        case Main(p) => TC(p, emptyenv)
    }
}
```

Output:
```
defined type alias Environment
TC: (input: Expr, gamma: Environment)Type
typecheck: (program: Program)Type
```

<!-- grade_id: cell-d9b206afd7a24388, points: 20, grade: true, solution: false -->
```scala
val p1 : Program = Main( Const(10) )
assert( typecheck(p1) == NumType, "test 1")
passed(20)
println("Hidden Tests : up to 80 points remaining")
```

Output:
```
*** Tests Passed (20 points) ***
Hidden Tests : up to 80 points remaining
null
```

<!-- grade_id: cell-31198cd330553090, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-69b68741c9fda826, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-d15f5cb6ed1607da, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-d8ea3308b64e150d, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-83408b1f20330fbc, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-8f34165e14575453, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-bb4cebe2260d7612, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-c4f5dcf0bf8fdca3, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-8037f9727ae33dd4, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-b59243a5817202d2, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-ede2ed89a8f56746, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-90690da2df5d6da5, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-288a52329f6c5428, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-7d0654077dff6d8b, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-d0489819cc1d8319, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-533d45c0493a5630, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-97f81d04838e2546, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-f91107fcf26e2240, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-3ba6c8052c148fff, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-36a0b90b6303d35e, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-89108e7186bf2e7e, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-e0c1267b4358e17e, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-7755f725e7af3e5c, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-a16534250c0ac8f9, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-6089a04546621af9, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-ea1c6a51893fd4d1, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-b3233ea6f1403203, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-0934c6879471b3a3, points: 2, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-750499f6519c9cb8, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-b9646dcc63b28149, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-3f118b29e49adbec, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-150fee1089cc254b, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```

<!-- grade_id: cell-f7fad77d24bd342e, points: 3, grade: true, solution: false -->
```scala
println("Hidden test")
```

Output:
```
Hidden test
```